# Question 2

More loops, more functions, and this time with conditionals!

### Part A

Write a function that tests if a number is even.

*HINT*: Recall the modulo operator `%`. This performs division on two numbers, but instead of returning their quotient, it returns the *remainder* of the division. For example, `7 % 3` returns 1, since the quotient is 2 with a remainder of 1.

Your function should:

 - be named `is_even`
 - take 1 argument, a number
 - return `True` if the argument is even, `False` otherwise
 
Make sure you're returning a *boolean*!

<!-- grade_id: q2a, points: 0, grade: false, solution: true -->
```python
def is_even(num):
    if (num %2) == 0:
        return True
    else:
        return False
    
```

<!-- grade_id: q2a_test1, points: 2.5, grade: true, solution: false -->
```python
assert is_even(3) == False
```

<!-- grade_id: q2a_test2, points: 2.5, grade: true, solution: false -->
```python
assert is_even(123456789) == False
```

<!-- grade_id: q2a_test3, points: 2.5, grade: true, solution: false -->
```python
assert is_even(42) == True
```

<!-- grade_id: q2a_test4, points: 2.5, grade: true, solution: false -->
```python
assert is_even(42.1) == False
```

### Part B

How could you test if a number is *odd*, using your `is_even` function from Part A?

Write your answer, either using 1 line of code or explaining in only 1 sentence, in the box below (double click on the statement below to activate the box).

<!-- grade_id: q2b, points: 5, grade: true, solution: true -->
I just need to change "equal to zero" to "not equal to zero" in the code. 

### Part C

Write a function that finds the *second-smallest* number in a list of numbers. Your function should:

 - be named `second_smallest`
 - take 1 argument, a list of numbers
 - return 1 number, the second-smallest number in the list
 - throw an `Exception` if there are not enough elements in the list
 
Take note of that last point. There's a possibility, after all, that the list your function receives as input will have 1 or even 0 elements in it. How, then, do you return the second-smallest element of a 0-element list? Well, simply, you can't. So instead, you'll raise an error.

Also note: you're the one *raising* the error, so it's *not* your job to catch it! It's the job of the person *using* your function to note the possibility that it will throw an error, and to take appropriate action to handle it gracefully. So if your code includes any `try` or `except` statements, you're doing it wrong!

For example, `second_smallest([1, 2, 3])` should return `2`. Another example `second_smallest([10, 100, 1000, 10000])` should return 100. And `second_smallest([1])` should raise an `Exception`.

You cannot use any built-in sorting functions! Only `range()` and `len()`.

<!-- grade_id: q2c, points: 0, grade: false, solution: true -->
```python
def second_smallest(x):
    if len(x)>=2:
        max1=max(x[0],x[1])
        max2=min(x[0],x[1])
        for i in range(2,len(x)):
            if x[i]>max1:
                max2=max1
                max1=x[i]
            else:
                if x[i]>max2:
                    max2=x[i]
    else:
        raise IndexError("Need more elements")
    return max2
```

```python
def second_smallest(x):
    max1=max(x[0],x[1])
    max2=min(x[0],x[1])
    for i in range(2,len(x)):
        if x[i]>max1:
            max2=max1
            max1=x[i]
        else:
            if x[i]>max2:
                max2=x[i]
    return max2
```

```python
x=[-100,100,100]
second_smallest(x)
```

Output:
```
100
```

```python
y=[]
second_smallest(y)
```

Output:
```
<ipython-input-6-39c97d36151b> in <module>
      1 y=[]
----> 2 second_smallest(y)

<ipython-input-4-99725b1f009a> in second_smallest(x)
      1 def second_smallest(x):
----> 2     max1=max(x[0],x[1])
      3     max2=min(x[0],x[1])
      4     for i in range(2,len(x)):
      5         if x[i]>max1:

IndexError: list index out of range
```

<!-- grade_id: q2c_test1, points: 5, grade: true, solution: false -->
```python
import numpy as np
np.random.seed(89547)

l1 = np.random.randint(-100, 100, 100)
a1 = np.sort(l1)[1]
np.testing.assert_allclose(a1, second_smallest(l1.tolist()))
```

Output:
```
~\AppData\Local\Continuum\anaconda3\lib\site-packages\numpy\testing\_private\utils.py in assert_allclose(actual, desired, rtol, atol, equal_nan, err_msg, verbose)
   1501     header = 'Not equal to tolerance rtol=%g, atol=%g' % (rtol, atol)
   1502     assert_array_compare(compare, actual, desired, err_msg=str(err_msg),
-> 1503                          verbose=verbose, header=header, equal_nan=equal_nan)
   1504 
   1505 

~\AppData\Local\Continuum\anaconda3\lib\site-packages\numpy\testing\_private\utils.py in assert_array_compare(comparison, x, y, err_msg, verbose, header, precision, equal_nan, equal_inf)
    827                                 verbose=verbose, header=header,
    828                                 names=('x', 'y'), precision=precision)
--> 829             raise AssertionError(msg)
    830     except ValueError:
    831         import traceback

AssertionError: 
Not equal to tolerance rtol=1e-07, atol=0

Mismatch: 100%
Max absolute difference: 191
Max relative differe
```

<!-- grade_id: q2c_test2, points: 5, grade: true, solution: false -->
```python
np.random.seed(485)

l2 = np.random.randint(0, 1000, 100)
a2 = np.sort(l2)[1]
np.testing.assert_allclose(a2, second_smallest(l2.tolist()))
```

<!-- grade_id: q2c_test3, points: 2.5, grade: true, solution: false -->
```python
try:
    second_smallest([])
except:
    assert True
else:
    assert False
```

<!-- grade_id: q2c_test4, points: 2.5, grade: true, solution: false -->
```python
try:
    second_smallest([100])
except:
    assert True
else:
    assert False
```

```python

```